package if_statements

// Advanced if statements and logical operators! ðŸ¹

// I AM NOT DONE

func CanVote(age int) string {
	// TODO: Use if with initialization: if age := checkAge(); condition
	// Return "Can vote" if age >= 18, "Cannot vote" otherwise
	// Fix me! Use if statement with proper logic
	return ""
}

func IsWeekend(day string) bool {
	// TODO: Return true if day is "Saturday" or "Sunday"
	// Use logical OR operator (||)
	// Fix me! Use logical operators
	return false
}

func IsValidPassword(password string) bool {
	// TODO: Return true if password meets ALL criteria:
	// - Length >= 8
	// - Contains at least one digit
	// Use logical AND operator (&&) and len() function
	// Hint: You can check for digits using a simple approach

	// Fix me! Implement password validation
	return false
}

func GetTemperatureAdvice(temp int) string {
	// TODO: Return advice based on temperature:
	// temp < 0: "Stay inside!"
	// 0 <= temp < 15: "Wear a coat"
	// 15 <= temp < 25: "Perfect weather"
	// temp >= 25: "Stay hydrated"
	// Fix me! Use nested if statements
	return ""
}
