package intro

// TODO: Add the missing import statement for the fmt package
// The fmt package is needed for the Printf function
// import "fmt"

// I AM NOT DONE

// PrintGreeting prints a greeting message to stdout
func PrintGreeting(name string) {
	// TODO: Fix this function to properly print a greeting
	// Hint: You need to import the fmt package first
	// Then use fmt.Printf with format "Hello, %s! Welcome to Go!\n"
	// For now, this function does nothing (but compiles)
}
