package intro

// Welcome to Whack-a-Goph! 🐹
// Your first gopher has popped up!
//
// This is the most basic exercise - make the Hello function
// return the string "Hello, Gopher!" to whack your first gopher!
//
// I AM NOT DONE

// TODO: Make this function return "Hello, Gopher!"
func Hello() string {
	// Fix me!
	return ""
}
