package variables

// Time to learn about variables! 🐹
// In Go, there are different ways to declare variables.
// Let's start with the most basic form.

// I AM NOT DONE

func Variables1() {
	// TODO: Declare a variable named 'x' with the value 5
	// Hint: var x int = 5
	x := "FIX ME!" // Remove this line and declare x properly above

	// TODO: Declare a variable named 'name' with the value "Gopher"
	// Hint: var name string = "Gopher"
	name := "FIX ME!" // Remove this line and declare name properly above

	// TODO: Declare a variable named 'done' with the value true
	// Hint: var done bool = true
	done := "FIX ME!" // Remove this line and declare done properly above

	// Don't modify anything below this line
	println("x =", x)
	println("name =", name)
	println("done =", done)
}
