package variables

// Now let's learn about type inference! 🐹
// Go can often figure out the type of a variable automatically.

// I AM NOT DONE

func Variables2() {
	// TODO: Declare a variable 'number' with value 42 using type inference
	// Hint: var number = 42
	number := "FIX ME!"

	// TODO: Declare a variable 'pi' with value 3.14 using type inference
	// Hint: var pi = 3.14
	pi := "FIX ME!"

	// TODO: Declare a variable 'message' with value "Hello, Go!" using type inference
	// Hint: var message = "Hello, Go!"
	message := "FIX ME!"

	// Don't modify anything below this line
	println("number is", number, "and has type int")
	println("pi is", pi, "and has type float64")
	println("message is", message, "and has type string")
}
