package primitive_types

// Time to learn about Go's primitive types! 🐹
// Go has several built-in basic types like int, string, bool, float64, etc.

// I AM NOT DONE

func GetAge() int {
	// TODO: Return your age as an integer
	// Fix me! This should return an actual number
	return 0
}

func GetName() string {
	// TODO: Return your name as a string
	// Fix me! This should return a non-empty string
	return ""
}

func IsStudent() bool {
	// TODO: Return true if you're a student, false otherwise
	// Fix me! This should return true or false
	return false
}

func GetPI() float64 {
	// TODO: Return the value of PI (approximately 3.14159)
	// Fix me! This should return the correct value
	return 0.0
}
