package primitive_types

// Let's explore more primitive types and constants! 🐹

// I AM NOT DONE

// TODO: Declare a constant named MaxInt with value 2147483647
const MaxInt = 0 // Fix me!

// TODO: Declare a constant named AppName with value "WhackAGoph"
const AppName = "" // Fix me!

func GetByte() byte {
	// TODO: Return the byte value 65 (which represents 'A' in ASCII)
	// Fix me!
	return 0
}

func GetRune() rune {
	// TODO: Return the rune value '🐹' (gopher emoji)
	// Fix me!
	return 0
}

func GetFloat32() float32 {
	// TODO: Return 1.5 as a float32
	// Fix me!
	return 0.0
}
