package main

// Tests 1: Basic Unit Testing 🐹
// Learn how to write unit tests using the testing package.

// I AM NOT DONE

// TODO: Implement Add function
func Add(a, b int) int {
	// Fix me!
	return 0
}

// TODO: Implement Multiply function
func Multiply(a, b int) int {
	// Fix me!
	return 0
}

// TODO: Implement IsEven function
func IsEven(n int) bool {
	// Fix me!
	return false
}

func main() {
	// This file focuses on testing, main is just for compilation
}
