package strings

// Advanced string operations with the strings package! 🐹

// I AM NOT DONE

func ToUpperCase(s string) string {
	// TODO: Convert string to uppercase
	// Hint: Use strings.ToUpper()
	// Fix me!
	return ""
}

func SplitWords(s string) []string {
	// TODO: Split string into words by spaces
	// Hint: Use strings.Split() or strings.Fields()
	// Fix me!
	return nil
}

func ReplaceAll(s, old, new string) string {
	// TODO: Replace all occurrences of 'old' with 'new' in string s
	// Hint: Use strings.ReplaceAll()
	// Fix me!
	return ""
}

func TrimSpaces(s string) string {
	// TODO: Remove leading and trailing whitespace
	// Hint: Use strings.TrimSpace()
	// Fix me!
	return ""
}

func HasPrefix(s, prefix string) bool {
	// TODO: Check if string s starts with prefix
	// Hint: Use strings.HasPrefix()
	// Fix me!
	return false
}
