package strings

// Time to learn about strings! 🐹
// Strings in Go are immutable and UTF-8 encoded.

// I AM NOT DONE

func GetLength(s string) int {
	// TODO: Return the length of string s
	// Hint: Use the len() function
	// Fix me!
	return 0
}

func Concatenate(a, b string) string {
	// TODO: Concatenate two strings and return the result
	// Hint: Use the + operator
	// Fix me!
	return ""
}

func GetFirstChar(s string) byte {
	// TODO: Return the first character of string s as a byte
	// Hint: Use indexing s[0], but make sure string is not empty
	// Fix me!
	return 0
}

func Contains(s, substr string) bool {
	// TODO: Check if string s contains substr
	// For now, just implement a simple check for exact match
	// Fix me!
	return false
}
