package vecs

// Time to learn about slices! 🐹
// Slices are Go's dynamic arrays, similar to Rust's Vec.

// I AM NOT DONE

func CreateSlice() []int {
	// TODO: Create and return a slice with values [1, 2, 3, 4, 5]
	// Fix me!
	return nil
}

func GetLength(slice []int) int {
	// TODO: Return the length of the slice
	// Hint: Use len(slice)
	// Fix me!
	return 0
}

func GetFirst(slice []int) int {
	// TODO: Return the first element of the slice
	// Assume slice is not empty
	// Fix me!
	return 0
}

func GetLast(slice []int) int {
	// TODO: Return the last element of the slice
	// Hint: Use slice[len(slice)-1]
	// Fix me!
	return 0
}

func AppendElement(slice []int, element int) []int {
	// TODO: Append element to slice and return the new slice
	// Hint: Use append(slice, element)
	// Fix me!
	return nil
}
