package conversions

// Time to learn about type conversions! 🐹
// Go requires explicit conversions between different types.

// I AM NOT DONE

func IntToFloat(i int) float64 {
	// TODO: Convert int to float64
	// Hint: Use float64(i)
	// Fix me!
	return 0.0
}

func FloatToInt(f float64) int {
	// TODO: Convert float64 to int (truncates decimal part)
	// Hint: Use int(f)
	// Fix me!
	return 0
}

func StringToInt(s string) int {
	// TODO: Convert string to int
	// For this exercise, assume the string is always a valid number
	// You can use a simple approach or return a fixed value for now
	// Fix me!
	return 0
}

func IntToString(i int) string {
	// TODO: Convert int to string
	// For this exercise, you can use a simple approach
	// Fix me!
	return ""
}

func ByteToRune(b byte) rune {
	// TODO: Convert byte to rune
	// Hint: Use rune(b)
	// Fix me!
	return 0
}
