package error_handling

// Time to learn about error handling! 🐹
// Go uses explicit error handling with the error type.

// I AM NOT DONE

// TODO: Implement the Divide function that handles division by zero
// The function should return the result of dividing a by b
// If b is 0, return an error with the message "division by zero"
// Otherwise, return the result and nil error
//
// Example: Divide(10, 2) should return (5.0, nil)
// Example: Divide(10, 0) should return (0.0, error)
func Divide(a, b float64) (float64, error) {
	// TODO: Check if b is 0
	// TODO: If b is 0, return 0 and errors.New("division by zero")
	// TODO: Otherwise, return a/b and nil
	return 0, nil
}
