package hashmaps

// Time to learn about maps! 🐹
// Maps are Go's key-value data structure, similar to Rust's HashMap.

// I AM NOT DONE

func CreateMap() map[string]int {
	// TODO: Create and return a map with keys "apple", "banana", "orange"
	// and values 5, 3, 8 respectively
	// Fix me!
	return nil
}

func GetValue(m map[string]int, key string) int {
	// TODO: Get the value for the given key from the map
	// Return 0 if key doesn't exist
	// Fix me!
	return 0
}

func SetValue(m map[string]int, key string, value int) {
	// TODO: Set the value for the given key in the map
	// Fix me!
}

func DeleteKey(m map[string]int, key string) {
	// TODO: Delete the key from the map
	// Hint: Use delete(m, key)
	// Fix me!
}

func GetKeys(m map[string]int) []string {
	// TODO: Return all keys from the map as a slice
	// Fix me!
	return nil
}
