package functions

// Time to learn about functions! 🐹
// Functions are the building blocks of Go programs.

// I AM NOT DONE

// TODO: Write a function called Add that takes two int parameters (a, b)
// and returns their sum as an int
// Example: Add(3, 4) should return 7
func Add(a, b int) int {
	// Fix me! This should return the sum of a and b
	return 0
}

// TODO: Write a function called Greet that takes a name (string) parameter
// and returns a greeting message: "Hello, [name]!"
// Example: Greet("Alice") should return "Hello, Alice!"
func Greet(name string) string {
	// Fix me! This should return a proper greeting
	return ""
}

// TODO: Write a function called IsEven that takes an int parameter
// and returns true if the number is even, false otherwise
// Hint: Use the modulo operator % to check if a number is divisible by 2
func IsEven(n int) bool {
	// Fix me! This should check if n is even
	return false
}
