package structs

// Advanced structs with methods! 🐹

// I AM NOT DONE

// TODO: Define a struct called Rectangle with fields Width and Height (both float64)
type Rectangle struct {
	Width  float64 // Fix me! This field exists but might need adjustment
	Height float64 // Fix me! This field exists but might need adjustment
}

// TODO: Add a method Area() to Rectangle that returns the area (Width * Height)
func (r Rectangle) Area() float64 {
	// Fix me!
	return 0.0
}

// TODO: Add a method Perimeter() to Rectangle that returns the perimeter
func (r Rectangle) Perimeter() float64 {
	// Fix me!
	return 0.0
}

// TODO: Add a method String() to Rectangle that returns a formatted string
// Format: "Rectangle: Width=X, Height=Y"
func (r Rectangle) String() string {
	// Fix me!
	return ""
}

// TODO: Define a struct called Circle with field Radius (float64)
type Circle struct {
	Radius float64 // Fix me! This field exists but might need adjustment
}

// TODO: Add a method Area() to Circle that returns the area (π * r²)
// Use 3.14159 for π
func (c Circle) Area() float64 {
	// Fix me!
	return 0.0
}

func CreateRectangle(width, height float64) Rectangle {
	// TODO: Create and return a Rectangle
	// Fix me!
	return Rectangle{}
}
